<br>
<div align="right">Enseignant : Aric Wizenberg</div>
<div align="right">E-mail : icarwiz@yahoo.fr</div>
<div align="right">Année : 2018/2019</div><br><br><br>
<div align="center"><span style="font-family:Lucida Caligraphy;font-size:32px;color:darkgreen">Master 2 MASERATI - Cours de Python</span></div><br><br>
<div align="center"><span style="font-family:Lucida Caligraphy;font-size:24px;color:#e60000">Correction des exercices du chapitre 2</span></div><br><br>
<hr>

Soit les constantes suivantes :

```python
CHEMIN_FICHIER = '../../../PyParis_XII_M1/cours/B.2.2.1_Nombres.ipynb'
SEPARATEUR_DE_CELLULES = '\n\n###\n\n'
```

```python
CHEMIN_FICHIER = 'B.2.2.1_Nombres.ipynb'
```

### Question
Charger le contenu du fichier **B.2.2.1_Nombres.ipynb** dont la localisation est contenue dans **CHEMIN_FICHIER** dans la variable **json_obj**, en utilisant le module **json**

<!-- grade_id: cell-10fefe2757582032, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
import json

with open(CHEMIN_FICHIER, encoding='utf-8') as fichier:
    contenu = fichier.read()

json_obj = json.loads(contenu)
### END SOLUTION
```

<!-- grade_id: cell-10fefe2757582032, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
import json

with open(CHEMIN_FICHIER, encoding='utf-8') as fichier:
    json_obj = json.load(fichier)
### END SOLUTION
```

<!-- grade_id: cell-a6ed0588ea6d4b5d, points: 1, grade: true, solution: false -->
```python
assert 'nbformat' in json_obj.keys(), 'Le contenu n\'a pas été bien chargé'
json_obj.keys()
```

```python
json_obj['cells'][0]['source']
```

### Question
Quel est le contenu de la première cellule ?

<!-- grade_id: cell-51c6b84b0a1c493b, points: 0, grade: false, solution: true -->
```python
reponse = None
### BEGIN SOLUTION
reponse = json_obj['cells'][0]
### END SOLUTION
```

<!-- grade_id: cell-1ead3cfb147b65cd, points: 1, grade: true, solution: false -->
```python
assert reponse['cell_type'] == 'markdown'
reponse['cell_type']
```

```python
reponse['source']
```

### Question
Créer une fonction **formattage_contenu** qui transforme un contenu du type de celui affiché comme résultat de la cellule précédente en une chaine unique de caractères. 

Cette chaine de caractères doit avoir pour derniers caractères la séquence contenue dans **SEPARATEUR_DE_CELLULES**.

<!-- grade_id: cell-21625a426866580c, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def formattage_contenu(source):        
    contenu = ''.join(source) + SEPARATEUR_DE_CELLULES  
    # print(contenu)
    return contenu
### END SOLUTION
```

<!-- grade_id: cell-68ad08cca918cba6, points: 1, grade: true, solution: false -->
```python
assert type(formattage_contenu(reponse['source'])) is str
assert formattage_contenu(reponse['source'])[-7:] == SEPARATEUR_DE_CELLULES
assert len(formattage_contenu(reponse['source'])) == 483

formattage_contenu(reponse['source'])
```

### Question
Utiliser une boucle pour définir deux listes **parties_de_code** et **parties_de_contenu**, à partir de la liste des cellules du Notebook contenue dans **json_obj**.

La première remplie avec les contenus de type **code** transformés grâce à la fonction **formattage_contenu**.

La seconde remplie avec les contenus de type **markdown** transformés grâce à la fonction **formattage_contenu**.

Le type d'une cellule est contenu dans l'entrée **cell_type**.

<!-- grade_id: cell-432ebbcc370595f6, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
parties_de_code = []
parties_de_contenu = []

for cellule in json_obj['cells']:
    if cellule['cell_type'] == 'markdown':
        parties_de_contenu.append(formattage_contenu(cellule['source']))
    elif cellule['cell_type'] == 'code':
        parties_de_code.append(formattage_contenu(cellule['source']))
### END SOLUTION
```

<!-- grade_id: cell-47cf6647eb9e3347, points: 1, grade: true, solution: false -->
```python
assert len(parties_de_code) == 21
assert len(parties_de_contenu) == 37
assert type(parties_de_code[0]) is str
assert type(parties_de_contenu[0]) is str
assert formattage_contenu(parties_de_code[0])[-7:] == SEPARATEUR_DE_CELLULES
assert formattage_contenu(parties_de_contenu[0])[-7:] == SEPARATEUR_DE_CELLULES

parties_de_code[0]
```

### Question

Ecrire dans le fichier **resultat.py** le contenu de parties_de_code

Ecrire dans le fichier **resultat.txt** le contenu de parties_de_contenu

<!-- grade_id: cell-275e76356206a47b, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
with open('resultat.py', 'w', encoding='utf-8') as destfile:
    destfile.writelines(parties_de_code)
    
with open('resultat.txt', 'w', encoding='utf-8') as destfile:
    for line in parties_de_contenu:
        destfile.write(line)
    
### END SOLUTION
```

<!-- grade_id: cell-2408baf2962a5aba, points: 1, grade: true, solution: false -->
```python
import os
assert os.path.getsize('resultat.py') == 479
assert os.path.getsize('resultat.txt') == 6975
```

# Formats de données itératifs : tables de données

## Exercice

### Intitalisation 

#### Question

Charger les bibliothèques nécessaires pour lire des tables de données (**numpy** et **pandas**)  et afficher des graphiques (**pyplot** et **seaborn**) en utilisant leurs alias usuels

<!-- grade_id: cell-2373cce6fe819907, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
import seaborn as sns
### END SOLUTION
```

On vérifie que tout a bien été chargé :

<!-- grade_id: cell-e0ae98daa2b6e791, points: 1, grade: true, solution: false -->
```python
print(np.__name__)
print(pd.__name__)
print(plt.__name__)
print(sns.__name__)
```

Mettez le format d'affichage graphique sur **inline**, et définissez la taille de figure par défaut sur **(12, 9)**.

<!-- grade_id: cell-21709098cdefa1e6, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
%matplotlib inline
plt.rcParams['figure.figsize'] = (12, 9) 
### END SOLUTION
```

<!-- grade_id: cell-fc52750ee1da50ca, points: 1, grade: true, solution: false -->
```python
# Cellule vide
### BEGIN HIDDEN TESTS
assert plt.rcParams['figure.figsize'] == [12, 9]
assert '%matplotlib inline' in In[-1]
### END HIDDEN TESTS
```

#### Question
Soit le chemin de fichier suivant :

```python
CHEMIN_DESC = '../../../PyParis_XII_M1/data/bike_share.txt'
```

Charger le texte contenu dans le fichier **bike_share.txt** dans la variable **metadonnees**

<!-- grade_id: cell-cc784900e7a87597, points: 0, grade: false, solution: true -->
```python
metadonnees = None
### BEGIN SOLUTION
with open(CHEMIN_DESC, encoding='utf-8') as fichier:
    metadonnees = fichier.read()
### END SOLUTION
```

<!-- grade_id: cell-454b06c335bfff06, points: 1, grade: true, solution: false -->
```python
assert len(metadonnees) == 664, 'Le fichier n\'a pas été bien chargé'
assert 'é' in metadonnees, 'Problème d\'encodage... Il fallait charger le fichier en utf-8'
print(metadonnees)
```

#### Question

Soit le chemin de fichier suivant :

```python
CHEMIN_DATA = '../../../PyParis_XII_M1/data/bike_share.csv'
```

Lire le contenu du fichier **bike_share.csv** et le mettre dans une variable **df_bike**

<!-- grade_id: cell-ef93d338fa6a2a52, points: 0, grade: false, solution: true -->
```python
df_bike = None
### BEGIN SOLUTION
df_bike = pd.read_csv(CHEMIN_DATA)
### END SOLUTION
```

<!-- grade_id: cell-f325921be49453d6, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert df_bike.shape == (10886, 12)
### END HIDDEN TESTS
df_bike
```

### Première analyse

#### Question

Quel est le nombre de ligne de cette table (ne pas écrire directement ce nombre bien sûr).

<!-- grade_id: cell-50f18d52e005977b, points: 0, grade: false, solution: true -->
```python
reponse = None
### BEGIN SOLUTION
reponse = df_bike.shape[0]
### END SOLUTION
```

<!-- grade_id: cell-7d5c79a7a23554e9, points: 1, grade: true, solution: false -->
```python
reponse
```

#### Question

Afficher les 5 dernières lignes du DataFrame

<!-- grade_id: cell-7ad3298ed9b48c3e, points: 0, grade: false, solution: true -->
```python
reponse = None
### BEGIN SOLUTION
reponse = df_bike.tail()
### END SOLUTION
```

<!-- grade_id: cell-0781be859de3a935, points: 1, grade: true, solution: false -->
```python
reponse
```

#### Question

Afficher les infos principales du DataFrame

<!-- grade_id: cell-089422806af7a452, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
df_bike.info()
### END SOLUTION
```

<!-- grade_id: cell-0907ed200e76de4a, points: 1, grade: true, solution: false -->
```python
# Cellule vide
```

#### Question
Afficher un extrait du DataFrame contenant toutes les lignes où les deux conditions suivantes sont respectées (voir les métadonnées pour trouver les identifiants correspondants) :
- le **weather** est de la catégorie : **Pluie ou neige légère** 
- la **season** est : **Printemps**

<!-- grade_id: cell-d3e2b8ca9e6de6a7, points: 0, grade: false, solution: true -->
```python
reponse = None
### BEGIN SOLUTION
reponse = df_bike[
    (df_bike['season'] == 1) & 
    (df_bike['weather'] == 3)
]
### END SOLUTION
```

<!-- grade_id: cell-f326eebc69c9e622, points: 1, grade: true, solution: false -->
```python
assert reponse.shape == (211, 12)
reponse
```

### Traitement des dates

#### Question

Créer une nouvelles série **date** dans **df_bike** qui contient la date dans un format compréhensible par Python (Timestamp/datetime)

```python
df_bike.loc[0, 'datetime']
```

<!-- grade_id: cell-4ee21ac9ead4635f, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
df_bike['date'] = pd.to_datetime(df_bike['datetime'])
### END SOLUTION
```

<!-- grade_id: cell-0383e3c930a85611, points: 1, grade: true, solution: false -->
```python
assert str(df_bike['date'].dtype) == 'datetime64[ns]'
df_bike['date'].sample(5)
```

#### Question
Créer une nouvelles série **heure** dans **df_bike** qui ne contient que l'heure en nombre entier (pas les minutes, juste l'heure, par exemple pour 16h15, on aurait juste 16), créez pour cela une fonction au préalable

```python
'2011-05-17 19:00:00'.split(' ')
```

```python
'19:00:00'.split(':')
```

<!-- grade_id: cell-2f211745f6e5d42d, points: 0, grade: false, solution: true -->
```python
def extraire_heure():
    pass

### BEGIN SOLUTION
def extraire_heure(x):
    partie_heure = x.split(' ')[1]
    heure_texte = partie_heure.split(':')[0]
    return int(heure_texte)

df_bike['heure'] = df_bike['datetime'].apply(extraire_heure)
### END SOLUTION
```

<!-- grade_id: cell-6976f7e4a2baedec, points: 1, grade: true, solution: false -->
```python
assert extraire_heure('2011-10-07 23:00:00') == 23
df_bike['heure'].sample(5)
```

#### Question

Créer un dictionnaire associant à un identifiant numérique le nom des différentes saisons

Faites de même pour les conditions météo

Aidez-vous pour cela des listes définies ci-dessous

```python
LISTE_SAISONS = ['Printemps', 'Eté', 'Automne', 'Hiver']
LISTE_CONDITIONS = ['Temps clair ou un peu couvert', 'Nuageux', 'Pluie ou neige légère', 'Grosse pluies ou neige']
```

```python
list(enumerate(LISTE_SAISONS))
```

<!-- grade_id: cell-faf7497d5e9479b6, points: 0, grade: false, solution: true -->
```python
dict_saisons = None
dict_meteo = None
### BEGIN SOLUTION
dict_saisons = {}
for i, v in enumerate(LISTE_SAISONS):
    dict_saisons[i + 1] = v
    
dict_meteo = {}   
for i, v in enumerate(LISTE_CONDITIONS):
    dict_meteo[i + 1] = v
    
# dict_saisons = {i + 1: v for i, v in enumerate(LISTE_SAISONS)}
# dict_meteo = {i + 1: v for i, v in enumerate(LISTE_CONDITIONS)}
### END SOLUTION
```

<!-- grade_id: cell-71ff09f876dbfc99, points: 1, grade: true, solution: false -->
```python
assert dict_saisons[3] == 'Automne'
assert dict_meteo[2] == 'Nuageux'
dict_saisons
```

#### Question

Modifier les séries **season** et **weather** en utilisant les dictionnaires précédemment crées afin que ce soient leurs noms qui soient contenus dans ces séries, à la place de l'identifiant numérique

<!-- grade_id: cell-2c70beb84839550e, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
df_bike['season'] = df_bike['season'].map(dict_saisons)
df_bike['weather'] = df_bike['weather'].map(dict_meteo)
### END SOLUTION
```

<!-- grade_id: cell-afe419db5d2559a9, points: 1, grade: true, solution: false -->
```python
print(df_bike['season'].sample(5))
print()
print(df_bike['weather'].sample(5))
```

#### Question

Créer une nouvelle série **mois** dans **df_bike** en utilisant une méthode adaptée de l'opérateur Series.dt appliquée à la séries **date**

<!-- grade_id: cell-1bd8e3d774203769, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
df_bike['mois'] = df_bike['date'].dt.month
### END SOLUTION
```

<!-- grade_id: cell-54f12f13b96dc9b1, points: 1, grade: true, solution: false -->
```python
df_bike[['date', 'mois']].sample(5)
```

#### Question

Charger le fichier 'mois_annee.csv' dans un nouveau DataFrame **df_mois**

```python
CHEMIN_MOIS_AN = '../../../PyParis_XII_M1/data/mois_annee.csv'
```

<!-- grade_id: cell-fc4b9503f9bec0d1, points: 0, grade: false, solution: true -->
```python
df_mois = None
## BEGIN SOLUTION
df_mois = pd.read_csv(CHEMIN_MOIS_AN, sep=';')
### END SOLUTION
```

<!-- grade_id: cell-1cc3604270a59f38, points: 1, grade: true, solution: false -->
```python
df_mois
```

#### Question

Utiliser le Dataframe des mois et une jointure pour ajouter la colonne month au DataFrame

<!-- grade_id: cell-76fc23037f60f8cd, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
df_bike = pd.merge(df_bike, df_mois, left_on='mois', right_on='id_month')
### END SOLUTION
```

<!-- grade_id: cell-7eb25230e083a167, points: 1, grade: true, solution: false -->
```python
df_bike[['date', 'mois', 'month']].sample(5)
```

#### Question

Supprimer la colonne datetime

<!-- grade_id: cell-36a0de61a0415855, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
df_bike.drop('datetime', axis=1, inplace=True)
### END SOLUTION
```

<!-- grade_id: cell-fe5c35b23735bc8e, points: 1, grade: true, solution: false -->
```python
assert 'datetime' not in df_bike.columns
```

### Analyse numérique

#### Question

Calculer la moyenne et l'écart-type de la série et les stocker dans des variables

<!-- grade_id: cell-c40565e800922853, points: 0, grade: false, solution: true -->
```python
moy_count = None
ecart_count = None
### BEGIN SOLUTION
moy_count = df_bike['count'].mean()
ecart_count = df_bike['count'].std()
### END SOLUTION
```

<!-- grade_id: cell-30a42702ff851e6a, points: 1, grade: true, solution: false -->
```python
print('moyenne :', moy_count)
print('écart-type :', ecart_count)
```

#### Question

Créer un tableau croisé dynamique ayant en ligne les saisons, et les conditions météos, puis les jours de la semaine et  pour colonnes la moyenne et la somme des **count**

<!-- grade_id: cell-1c80afe00a07adb9, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
df_heu_sai = df_bike.pivot_table(
    values='count', 
    index=['season', 'weather'], 
    aggfunc=[np.mean, np.sum]
)
### END SOLUTION
```

<!-- grade_id: cell-5e89b9403fa99bed, points: 1, grade: true, solution: false -->
```python
assert df_heu_sai.shape == (13, 2)
df_heu_sai
```

#### Question
Afficher un décompte du nombre de périodes de chacun des types de météo répertoriés dans **df_bike**

<!-- grade_id: cell-d28d99e34058b8b3, points: 0, grade: false, solution: true -->
```python
reponse = None
### BEGIN SOLUTION
reponse = df_bike.weather.value_counts()
### END SOLUTION
```

<!-- grade_id: cell-b8b10eb900bdf000, points: 1, grade: true, solution: false -->
```python
assert reponse.loc['Temps clair ou un peu couvert'] == 7192
reponse
```

### Analyse graphique

#### Question

Utiliser la fonction boxplot de seaborn pour afficher la boîte à moustache pour la variable count

<!-- grade_id: cell-711789f4665ea172, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
sns.boxplot(df_bike['count'])
### END SOLUTION
```

<!-- grade_id: cell-76494b09afb233e0, points: 1, grade: true, solution: false -->
```python

```

#### Question

Utiliser maintenant cette même fonction avec :
- la série **heure** en tant que paramètre x pour décliner count selon l'heure de la journée
- la série **count** en tant que paramètre y
- l'option **orient='v'** pour renverser le graphique

<!-- grade_id: cell-404c24dd56855217, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
sns.boxplot(df_bike['heure'], df_bike['count'], orient='v')
### END SOLUTION
```

<!-- grade_id: cell-3f015d942bb1ea57, points: 1, grade: true, solution: false -->
```python

```

#### Question

Afficher une **heatmap** des **correlations** entre les variables suivantes :

'temp','atemp','casual','registered','humidity','windspeed','count'

Ajouter l'option **annot=True** pour afficher les valeurs sur le graphique.

<!-- grade_id: cell-9fd1c7580b674ea8, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
correlations = df_bike[['temp','atemp','casual','registered','humidity','windspeed','count']].corr()
sns.heatmap(correlations, annot=True)
### END SOLUTION
```

<!-- grade_id: cell-d1317271f3dde386, points: 1, grade: true, solution: false -->
```python

```

#### Question
Afficher le distplot de la série **count**

<!-- grade_id: cell-4262bce7795d5f88, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
sns.distplot(df_bike['count'])
### END SOLUTION
```

<!-- grade_id: cell-bd680e56ccb52163, points: 1, grade: true, solution: false -->
```python

```

```python

```